---
title: 'Test H2O in Python'
description: 'Chapter description goes here.'
---

## Initialize H2O

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}
import h2o
h2o.init()
```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
